## HW3: Kernel Ridge Regression

In this notebook you will review the details of kernel ridge regression (KRR). In particular:

* How to formulate matrices and solutions to KRR
* Use scikit-learn to train a KRR model for the problem in HW2

```python
import pandas as pd
import numpy as np
import seaborn as sns
import matplotlib.pyplot as plt
from sklearn.kernel_ridge import KernelRidge
from sklearn.metrics import mean_squared_error, mean_absolute_error
%matplotlib inline
```


## Data Loading

For the first part of this HW, you will predict fuel efficiency (`mpg`) of several models of automobiles using a **single feature**.

First, let's load in the data. We're using a built in dataset for the Seaborn library so there is no external download needed.

```python
# Here, we load the fuel dataset, and drop any rows that have missing data.
vehicle_data = sns.load_dataset('mpg').dropna()
vehicle_data = vehicle_data.sort_values('horsepower', ascending=True)
vehicle_data.head(5)
```

```python
vehicle_data.shape
```

We have 392 datapoints and 8 potential features (plus our observed $y$ values, `mpg`).

Typically for regression, we want our data is nice simple arrays so we can do something like fit(x,y). For that reason, it can be helpful to just define table columns as variables for later use

```python
v_mpg = vehicle_data['mpg']
v_horsepower = vehicle_data['horsepower']
```

Let's take a look at the data

```python
# Run this cell to visualize the data. 
sns.scatterplot(data=vehicle_data, x='horsepower', y='mpg');
plt.title("mpg vs horsepower");
```

## KRR with a Single Variable

It is evident that mpg does not depend on horsepower linearly, so you need to use nonlinear kernels for KRR. Let's review the key ideas behind KRR.

The KRR model reads

$$
\tilde{y} = \sum_{i=1}^m \alpha_i k(x_i, x),
$$

where $x_i$ is the input of the $i$th training data point, $m$ is the number of training data points, and $k(x_i, x$) is the kernal function. Like in linear regression, we would like to minimize the squared error with a ridge regularization, namely

$$
L(\alpha_i) = \sum_{i=1}^m (y_i - \tilde{y}_i)^2 + \lambda \sum_{i=1}^m \alpha_i^2,
$$
where $\lambda$ is the regularization parameter, another hyperparameter. If one sets the first derivatives of the loss function with respect to $\alpha_i$ to be zero, after some matrix algebra, one obtains the normal equation for KRR:

$$
\vec{\alpha} = ({\bf K} + \lambda {\bf I})^{-1} \vec{y},
$$
where $\vec{\alpha} = (\alpha_1, \alpha_2, \cdots, \alpha_m)^T$, $\vec{y} = (y_1, y_2, \cdots, y_m)^T$, $\bf I$ is the identity matrix, and the matrix element of the kernel matrix (a.k.a., Gram matrix) is ${\bf K}_{ij} = k(x_i, x_j)$. For the linear kernel, $\bf K = XX^T$. For other nonlinear kernels (e.g., Gaussian, polynomial, etc), one can compute $\bf K$ using its definition. 

```python
# Instruction: here you will write your own code to do
# Input: horsepower
# Output: mpg
# Kernel: your choice, and make sure you play around with hyperparameters to see if you can improve the prediction
# Final results: 
# 1. Plot actual and predicted mpg vs horsepower in the same plot
# 2. Compute Mean Squared Error (MSE) and Mean Average Error (MAE)
# 3. Compute the correlation coefficient, r
# 4. Make the parity plot, namely predicted mpg vs actual mpg. Make sure the parity plot is square.

# You are highly encouraged to use AI assistant to help you write the code.

# Reference: https://scikit-learn.org/stable/modules/generated/sklearn.kernel_ridge.KernelRidge.html
```

## KRR with Multiple Variables

Clearly, mpg does not just depend on horsepower. For example, rows 3-5 all have a horsepower of 48.0, but their mpg are different (43.1, 43.4, 44.3). Other factors must also play a role. Without looking into the data closely (one should in real research), one might ignore origin and name, but consider the other 6 features. Let's do KRR with these 6 features.

```python
# Instruction: here you will write your own code to do
# Input: horsepower, cylinders, displacement, weight, acceleration, model_year
# Output: mpg
# Kernel: your choice, and make sure you play around with hyperparameters to see if you can improve the prediction
# Final results: 
# 1. Plot actual and predicted mpg vs horsepower in the same plot
# 2. Compute Mean Squared Error (MSE) and Mean Average Error (MAE)
# 3. Compute the correlation coefficient, r
# 4. Make the parity plot, namely predicted mpg vs actual mpg. Make sure the parity plot is square.

# You are highly encouraged to use AI assistant to help you write the code.

# Reference: https://scikit-learn.org/stable/modules/generated/sklearn.kernel_ridge.KernelRidge.html
```

```python

```

Tests for q1a (points: None)
```
>>> X.shape == (392,2)
True
```
```
>>> (add_intercept(np.array([[1, 2, 3],[4, 5, 6]]).T)[:,0] == np.ones((3,))).all()
True
```
```
>>> add_intercept(np.array([[1, 2, 3],[4, 5, 6]]).T).shape == (3,3)
True
```
```
>>> (add_intercept(np.array([[1, 2, 3],[4, 5, 6]]).T)[:,2] == np.array([4,5,6])).all()
True
```


Tests for q1b (points: None)
```
>>> linear_model(np.arange(1,5), np.arange(1,5)) == 30
True
```
```
>>> (linear_model(2*np.eye(100), np.ones(100)) == 2*np.ones(100)).all()
True
```
```
>>> test_theta = np.array([[1, 2], [3, 4], [5, 6]])
>>> test_x = np.array([[1, 3, 5], [2, 4, 6]])
>>> expected = np.array([[35, 44], [44, 56]])
>>> actual = linear_model(test_theta, test_x)
>>> np.array_equal(actual, expected)
True
```
```
>>> test_theta_2 = np.array([[3], [5]])
>>> test_x_2 = np.array([[1, 4], [1, 6], [1, 8]])
>>> expected_2 = np.array([[23], [33], [43]])
>>> actual_2 = linear_model(test_theta_2, test_x_2)
>>> np.array_equal(expected_2, actual_2)
True
```


Tests for q1c (points: None)
```
>>> analytical_thetas.shape in ((2,), (2, 1))
True
```
```
>>> np.isclose(analytical_thetas[0], 39.93586102)
True
```


Tests for q3a (points: None)
```
>>> np.isclose(model_multi.intercept_, 56.900099702112954)
True
```
```
>>> np.isclose(model_multi.coef_[0], -0.46618963)
True
```
```
>>> np.isclose(model_multi.coef_[1], 0.00123054)
True
```
```
>>> # Make sure our old model is not overwritten.
>>> np.isclose(model.intercept_, 58.70517203721748) and np.isclose(model.coef_[0], -3.50352375)
True
```


Tests for q3c (points: None)
```
>>> np.isclose(mean_predicted_mpg_hp_only, np.mean(predicted_mpg_hp_only), atol=0.001)
True
```
```
>>> np.isclose(mean_predicted_mpg_hp_sqrt, np.mean(predicted_mpg_hp_sqrt), atol=0.001)
True
```
```
>>> np.isclose(mean_predicted_mpg_multi, np.mean(predicted_mpg_multi), atol=0.001)
True
```


Tests for q4a (points: None)
```
>>> X_redundant.shape == (392,4)
True
```
